<a href="https://colab.research.google.com/github/laguz/DS-Unit-4-Sprint-2-Neural-Networks/blob/dependabot%2Fpip%2Ftensorflow-2.5.0/LS_DS_Unit_4_Sprint_Challenge_2_AG.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>


## Autograded Notebook (Canvas & CodeGrade)

This notebook will be automatically graded. It is designed to test your answers and award points for the correct answers. Following the instructions for each Task carefully.
Instructions

- **Download** this notebook as you would any other ipynb file 
- **Upload** to Google Colab or work locally (if you have that set-up)
- **Delete** `raise NotImplementedError()`

- **Write** your code in the `# YOUR CODE HERE` space


- **Execute** the Test cells that contain assert statements - these help you check your work (others contain hidden tests that will be checked when you submit through Canvas)

- **Save** your notebook when you are finished
- **Download** as a ipynb file (if working in Colab)
- **Upload** your complete notebook to Canvas (there will be additional instructions in Slack and/or Canvas)




## *Data Science Unit 4 Sprint 2*

# Sprint Challenge - Neural Network Foundations

Table of Problems

1. [Defining Neural Networks](#Q1)
2. [Simple Perceptron](#Q2)
    - Perceptron
    - Multilayer Perceptron (i.e. Neural Network)
    - Analyze and Compare
4. [Keras MMP](#Q3)


____

# Before you submit your notebook you must first

1) Restart your notebook's Kernel

2) Run all cells sequentially, from top to bottom, so that cell numbers are sequential numbers (i.e. 1,2,3,4,5...)
- Easiest way to do this is to click on the **Cell** tab at the top of your notebook and select **Run All** from the drop down menu. 

____

<a id="Q1"></a>
## 1. Defining Neural Networks 

Write *your own* definitions for the following terms:

- **Neuron:** is a unit that take input data and give a output base on the data.

- **Input Layer:** is the first layer of neuron of a Neural Networks.

- **Hidden Layer:** is the layer/s of the Neural Networks that change the input layer data to pass to the output layer.

- **Output Layer:** is the layer that get the process data from the hidden layer/s

- **Activation:** is the function that determine the output of the Neural Networks.

<!-- grade_id: cell-6adae65226f09553, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

- `Explain` how Back-propagation works 
- `Explain` how Gradient Descent works (mention the learning rate)
- `Explain` how Back-propagation and Gradient Descent are related   

Use your own words, but feel free to reference external materials for this question. 

<!-- grade_id: cell-cef20b23d4e0b056, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

Remember our Simple Perceptron Class from Monday. 

- Describe the process of making a prediction, i.e. how do you go from inputs to outputs?

<!-- grade_id: cell-53c7cc36db9d7983, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

<a id="Q2"></a>
## 2. Simple Perceptron

In this question, you will build two neural networks using `Keras`. After you build these two models, compare the results of the model.

```python
"""
Our Dataset
"""
import numpy as np


rng = np.random.RandomState(0)

"Use this X & y in the following 2 models"
X = rng.randn(300, 2)
y = np.array(np.logical_xor(X[:, 0] > 0, X[:, 1] > 0), 
             dtype=int)
```

```python
X.shape
```

Output:
```
(300, 2)
```

```python
y
```

Output:
```
array([0, 0, 1, 1, 1, 0, 0, 0, 1, 1, 1, 1, 1, 1, 0, 0, 0, 1, 0, 0, 0, 1,
       0, 1, 0, 1, 0, 1, 0, 0, 0, 0, 1, 1, 1, 0, 1, 1, 0, 1, 1, 1, 0, 1,
       1, 1, 0, 0, 0, 0, 1, 1, 1, 0, 0, 0, 1, 1, 1, 0, 1, 0, 0, 1, 0, 1,
       1, 1, 1, 1, 1, 0, 0, 1, 0, 1, 0, 0, 1, 0, 1, 0, 1, 0, 0, 0, 1, 1,
       0, 1, 0, 0, 1, 1, 1, 1, 0, 1, 0, 0, 0, 0, 1, 0, 1, 0, 1, 0, 1, 1,
       0, 1, 1, 0, 0, 0, 1, 1, 1, 0, 0, 0, 0, 0, 0, 1, 1, 0, 0, 1, 0, 0,
       1, 0, 1, 1, 0, 0, 1, 1, 1, 1, 0, 1, 0, 0, 1, 1, 0, 1, 1, 0, 1, 0,
       1, 1, 0, 1, 0, 0, 0, 1, 0, 1, 1, 1, 1, 1, 1, 1, 1, 0, 0, 1, 0, 0,
       1, 0, 0, 1, 0, 0, 1, 0, 0, 1, 1, 0, 1, 1, 1, 0, 1, 0, 0, 1, 0, 0,
       0, 0, 0, 0, 1, 0, 0, 1, 0, 0, 0, 0, 0, 1, 1, 1, 1, 0, 0, 1, 0, 0,
       0, 1, 1, 1, 0, 1, 0, 0, 0, 1, 0, 1, 1, 0, 1, 0, 0, 0, 1, 0, 0, 0,
       1, 0, 1, 0, 0, 1, 0, 0, 1, 1, 0, 0, 1, 0, 0, 1, 0, 0, 0, 0, 1, 0,
       1, 1, 0, 0, 0, 0, 1, 1, 0, 1, 1, 0, 1, 1, 1, 0, 1, 1, 1, 0, 0, 1,
       0, 1, 0, 1, 0, 1, 0, 1, 0, 0, 0, 0, 0, 0])
```

```python
2**2
```

Output:
```
4
```

```python
4**4
```

Output:
```
256
```

```python
3**3
```

Output:
```
27
```

```python

```

This word is speled wrong. 

### Simple Perceptron
Construct a simple perceptron using Keras. 

Make sure to include the following in your model:
- Add `1 dense layer` with a `single neuron` 
- Use a `sigmoid activation function`
- Set `epochs` to 10 
- Use the version of `crossentropy loss` that is appropriate for this data.

Your model should be called `model1` and make sure to save the results of your fit statement to a variable called `h1`. 


<!-- grade_id: cell-427690628f9c900b, points: 0, grade: false, solution: true -->
```python
import tensorflow
from tensorflow.keras.models import Sequential
from tensorflow.keras.layers import Dense

# build and fit model

# YOUR CODE HERE
# Perceptron model

model1 = Sequential()
model1.add(Dense(1,input_dim=2, activation='sigmoid'))
model1.compile(loss='binary_crossentropy', optimizer='adam', metrics=['accuracy'])
model1.fit(X,y, epochs=10);
```

Output:
```
Epoch 1/10
10/10 [==============================] - 0s 1ms/step - loss: 0.7299 - accuracy: 0.4833
Epoch 2/10
10/10 [==============================] - 0s 1ms/step - loss: 0.7284 - accuracy: 0.4833
Epoch 3/10
10/10 [==============================] - 0s 1ms/step - loss: 0.7267 - accuracy: 0.4833
Epoch 4/10
10/10 [==============================] - 0s 1ms/step - loss: 0.7252 - accuracy: 0.4833
Epoch 5/10
10/10 [==============================] - 0s 1ms/step - loss: 0.7237 - accuracy: 0.4800
Epoch 6/10
10/10 [==============================] - 0s 1ms/step - loss: 0.7223 - accuracy: 0.4800
Epoch 7/10
10/10 [==============================] - 0s 1ms/step - loss: 0.7210 - accuracy: 0.4800
Epoch 8/10
10/10 [==============================] - 0s 1ms/step - loss: 0.7197 - accuracy: 0.4800
Epoch 9/10
10/10 [==============================] - 0s 1ms/step - loss: 0.7184 - accuracy: 0.4733
Epoch 10/10
10/10 [==============================] - 0s 1ms/step - loss: 0.7172 - accuracy: 0.4700
```

<!-- grade_id: cell-bf2ae566afacde8c, points: 1, grade: true, solution: false -->
```python
# Visible test
assert len(model1.get_config()["layers"]) == 2, "Make sure you only create 1 Dense layer."
assert len(h1.epoch) <=10, "Did you make sure to set epochs to 10 or less?"
```

Output:
```
AttributeError                            Traceback (most recent call last)
<ipython-input-16-4bb039868c41> in <module>()
      1 # Visible test
      2 assert len(model1.get_config()["layers"]) == 2, "Make sure you only create 1 Dense layer."
----> 3 assert len(h1.epoch) <=10, "Did you make sure to set epochs to 10 or less?"

AttributeError: 'int' object has no attribute 'epoch'
```

```python
model1.get_config()["layers"]
```

Output:
```
[{'class_name': 'InputLayer',
  'config': {'batch_input_shape': (None, 2),
   'dtype': 'float32',
   'name': 'dense_3_input',
   'ragged': False,
   'sparse': False}},
 {'class_name': 'Dense',
  'config': {'activation': 'sigmoid',
   'activity_regularizer': None,
   'batch_input_shape': (None, 2),
   'bias_constraint': None,
   'bias_initializer': {'class_name': 'Zeros', 'config': {}},
   'bias_regularizer': None,
   'dtype': 'float32',
   'kernel_constraint': None,
   'kernel_initializer': {'class_name': 'GlorotUniform',
    'config': {'seed': None}},
   'kernel_regularizer': None,
   'name': 'dense_3',
   'trainable': True,
   'units': 1,
   'use_bias': True}}]
```

<!-- grade_id: cell-a957e14380b2f508, points: 1, grade: true, solution: false -->
```python
# Hidden tests - you will see the results when you submit to Canvas
```

### Multi-Layer Perceptron
Now construct a multi-layer perceptron model (also known as a neural network). 

Your neural network `must` have: 
- `2` Hidden Layers
- Select any number between `5-32` for the number of neurons in each hidden layers
- Your pick of activation function and optimizer
- Incorporate the `Callback function` below into your model
- Set epochs to `100`
- Your model should be called `model2` 
- Save the results of your fit statement to a variable called `h2`. 
- Use the version of `crossentropy loss` that is appropriate for this data.

```python
import tensorflow as tf

class myCallback(tf.keras.callbacks.Callback): 
    def on_epoch_end(self, epoch, logs={}): 
        # if model reaches 99% accuracy, training is terminated 
        acc_threshold = 0.99
        if(logs.get('accuracy') > acc_threshold):   
            self.model.stop_training = True
            self.model.callback_used = True
```

<!-- grade_id: cell-77523c4c64743f16, points: 0, grade: false, solution: true -->
```python
# build and fit model

# YOUR CODE HERE
model2 = Sequential()
model2.add(Dense(10,input_dim=2, activation='sigmoid'))
model2.add(Dense(10, activation='relu'))
model2.add(Dense(10, activation='relu'))

model2.compile(loss='binary_crossentropy', optimizer='adam', metrics=['accuracy'])
model2.fit(X,y, epochs=100);
```

Output:
```
Epoch 1/100
10/10 [==============================] - 0s 1ms/step - loss: 3.7257 - accuracy: 0.4733
Epoch 2/100
10/10 [==============================] - 0s 2ms/step - loss: 3.5269 - accuracy: 0.4633
Epoch 3/100
10/10 [==============================] - 0s 2ms/step - loss: 3.4388 - accuracy: 0.4000
Epoch 4/100
10/10 [==============================] - 0s 2ms/step - loss: 3.4080 - accuracy: 0.3367
Epoch 5/100
10/10 [==============================] - 0s 1ms/step - loss: 3.3907 - accuracy: 0.2867
Epoch 6/100
10/10 [==============================] - 0s 2ms/step - loss: 3.3770 - accuracy: 0.2133
Epoch 7/100
10/10 [==============================] - 0s 2ms/step - loss: 3.3671 - accuracy: 0.1800
Epoch 8/100
10/10 [==============================] - 0s 2ms/step - loss: 3.3600 - accuracy: 0.1533
Epoch 9/100
10/10 [==============================] - 0s 1ms/step - loss: 3.3544 - accuracy: 0.1367
Epoch 10/100
10/10 [==============================] - 0s 1ms/step - loss: 3.3496 - accuracy: 0.1333
Epoch 11/
```

<!-- grade_id: cell-770612ca24334d8a, points: 1, grade: true, solution: false -->
```python
# Visible test
assert len(model2.get_config()["layers"]) == 4, "You should have 4 layers: Input, hidden 1, hidden 2, output."
assert 5 <= model2.get_config()["layers"][1]["config"]["units"] <= 32, "You should have 5 - 32 units in hidden layer 1, but don't."
assert 5 <= model2.get_config()["layers"][2]["config"]["units"] <= 32, "You should have 5 - 32 units in hidden layer 2, but don't."
assert h2.params["epochs"] == 100, "You didn't set epochs to 100."
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-57-769a996437d0> in <module>()
      3 assert 5 <= model2.get_config()["layers"][1]["config"]["units"] <= 32, "You should have 5 - 32 units in hidden layer 1, but don't."
      4 assert 5 <= model2.get_config()["layers"][2]["config"]["units"] <= 32, "You should have 5 - 32 units in hidden layer 2, but don't."
----> 5 assert h2.params["epochs"] == 100, "You didn't set epochs to 100."

NameError: name 'h2' is not defined
```

```python

```

Output:
```
AttributeError                            Traceback (most recent call last)
<ipython-input-59-e1b816fd35ad> in <module>()
----> 1 model1.params["epochs"]

AttributeError: 'Sequential' object has no attribute 'params'
```

<!-- grade_id: cell-49b1bf7cce22b5b9, points: 1, grade: true, solution: false -->
```python
# Hidden tests - you will see the results when you submit to Canvas
```

### Analyze and Compare

**Before you Start**: You will need to install an additional library for this next segment. 

Install the package `mlxtend` into the environment you are using for the sprint challenge.

You can install this package using the following statement in the terminal

```python
pip install mlxtend
```

Or you can install this package using the following statement in your notebook

```python
!pip install mlxtend
```

If you choose to install this package from within your notebook, be sure to delete the install statement afterwards so that CodeGrade doesn't try to install it and potentially crash. 


The cells below generate decision boundary plots of your models (`model1` & `model2`). Review the plots.

```python
X.shape, y.shape
```

```python
import matplotlib.pyplot as plt
from mlxtend.plotting import plot_decision_regions


fig = plt.figure(figsize=(12,6))

for clf, hist, name, grd in zip([model1,model2], [h1, h2],['Perceptron', 'Multi-Layer Perceptron'],[1,2]):

    ax = plt.subplot(1,2, grd)
    fig = plot_decision_regions(X=X, y=y, clf=clf, legend=2)
    title = f"{name} with {hist.history['accuracy'][-1]:,.2f} Accuracy"
    plt.title(title)

plt.show()
```

1) Why does the Perceptron (`model1`) only achieve about 50% accuracy? 

2) What is the architectural property of the Multi-Layer Perceptron that allows it to more accurately learn the relationship between X and y? 
- Hint: recall that each layer represents a vector space and they usually have a different number of dimensions, $\mathbb{R}^N$.
   

<!-- grade_id: cell-302694c508c8da0e, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## 3. Keras MMP <a id="Q3"></a>

- Implement a Multilayer Perceptron architecture of your choosing using the Keras library. 
- Train your model and report its baseline accuracy. 
- Then `hyper-parameters tune two parameters each with no more than 2 values each`
    - Due to limited computational resources on CodeGrade `DO NOT INCLUDE ADDITIONAL PARAMETERS OR VALUES PLEASE`
- Report your optimized model's accuracy
- Use the Heart Disease Dataset provided (binary classification)
- Use an appropriate loss function for a binary classification task
- Use an appropriate activation function on the final layer of your network.
- Train your model using verbose output for ease of grading.
- Use GridSearchCV to hyper-parameters tune your model. 
    - **Use `n_jobs` = 1**
- When hyper-parameters tuning, show you work by adding code cells for each new experiment.
- Report the accuracy for each combination of hyper-parameters as you test them so that we can easily see which resulted in the highest accuracy.

```python
import pandas as pd
from sklearn.preprocessing import StandardScaler

# load data
data_path = 'https://raw.githubusercontent.com/ryanleeallred/datasets/master/heart.csv'
df = pd.read_csv(data_path)
df = df.sample(frac=1)
print(df.shape)
df.head()
```

<!-- grade_id: cell-85dc40f19f5a1d6b, points: 0, grade: false, solution: true -->
```python
# Create an input matrix named 'X' store it in a 2D numpy array

# Create an output vector for the labels named 'Y', store it in 1D numpy array

# YOUR CODE HERE
raise NotImplementedError()
```

```python
X.shape
```

<!-- grade_id: cell-c17c686c974edc2e, points: 1, grade: true, solution: false -->
```python
# Visible Testing
assert X.shape[0] == 303, "Did you drop/lose some rows in X? Did you properly load and split the data?"
assert X.shape[1] == 13, "Did you drop/lose some columns in X? Did you properly load and split the data?"
assert len(Y)== 303, "Did you drop/lose some rows in Y? Did you properly load and split the data?"
```

```python
import tensorflow.keras
from tensorflow.keras.models import Sequential
from tensorflow.keras.layers import Dense
from sklearn.model_selection import GridSearchCV
from tensorflow.keras.wrappers.scikit_learn import KerasClassifier
```

<!-- grade_id: cell-cfc5517cd0b6fa64, points: 0, grade: false, solution: true -->
```python
# Create a function named 'create_model' that returns a complied keras model -  required for KerasClassifier
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-fac25126eaf1eee4, points: 1, grade: true, solution: false -->
```python
# Visible Testing
assert create_model().__module__ == 'tensorflow.python.keras.engine.sequential', "create_model should return a keras model that was created using the Sequential class."
```

<!-- grade_id: cell-fbc3d0a07230078c, points: 0, grade: false, solution: true -->
```python
# Pass 'create_model' into KerasClassifier, store KerasClassifier to a variable named 'model'
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-464e7506993775f2, points: 1, grade: true, solution: false -->
```python
# Visible Testing
assert model.__module__ == 'tensorflow.python.keras.wrappers.scikit_learn', "model should be a instance of KerasClassifier."
```

<!-- grade_id: cell-985c0425f3b1304d, points: 0, grade: false, solution: true -->
```python
# Define the grid search parameters inside a dictionary named 'param_grid' 
# Use 2 hyper-parameters with 2 possible values for each 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c765b5db5489d7a2, points: 1, grade: true, solution: false -->
```python
assert len(param_grid.keys()) == 2, "Did you create a param dict with 2 hyper-parameters as keys?"
```

<!-- grade_id: cell-7cfb4315eab5031c, points: 0, grade: false, solution: true -->
```python
# Create Grid Search object and name it 'gs'
# Run Grid Search 
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# your grid_result object should be able to run in this code 
print(f"Best: {grid_result.best_score_} using {grid_result.best_params_}")
means = grid_result.cv_results_['mean_test_score']
stds = grid_result.cv_results_['std_test_score']
params = grid_result.cv_results_['params']
for mean, stdev, param in zip(means, stds, params):
    print(f"Means: {mean}, Stdev: {stdev} with: {param}") 
```

```python

```